Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE".

Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All). On JupyterLab, you may want to hit the "Validate" button as well.

Caution: do not mess with the notebook's metadata; do not change a pre-existing cell's type; do not copy pre-existing cells (add new ones with the + button instead). This will break autograding; you will get a 0; you are warned.

<table style="width: 100%; border: none;" cellspacing="0" cellpadding="0" border="0">
  <tr>
    <td><img src="https://cat.opidor.fr/images/thumb/c/cf/CMJN.png/1200px-CMJN.png" style="float: left; width: 100%" />
</td>
    <td><a style="font-size: 3em; text-align: center; vertical-align: middle;" href="https://moodle.ip-paris.fr/course/view.php?id=10603">[CSC2S004EP - 2025] - Introduction to Machine Learning</a>
</td>
  </tr>
</table>

<a style="font-size: 3em;">Lab Session 1: Jupyter Notebooks, Pandas, Sklearn basics</a>

# Introduction

Welcome to the 1st [CSC2S004EP](https://moodle.ip-paris.fr/course/view.php?id=10603) lab session: *Introduction to Machine Learning*.

This lab session assumes you are familiar with the Python programming language.

Note that the 1st year courses used scripting/packaging tools (*e.g.* Spyder) and relied on relatively "low-level" abstractions (lists, numpy arrays, etc.) and manipulation techniques (*e.g.* `for` loops). While the knowledge of these concepts and tools is mandatory for this course and your future use of Python, we will concentrate in this course respectively on Jupyter Notebooks (more on these later) to write code and "high-level", "more Pythonic" abstractions (e.g. `pandas`, scikit-learn, plotting libraries) as a coding style.

Additional documentation can be found there:
- [Python 3 documentation](https://docs.python.org/3/)
- [The Python Tutorial](https://docs.python.org/3/tutorial/index.html)
- [The Python Standard Library](https://docs.python.org/3/library/index.html)
- [The Python Language Reference](https://docs.python.org/3/reference/index.html)

# Jupyter Notebooks

The document you are reading is a [Jupyter Notebook](https://jupyter.org/).

The Jupyter Notebook (formerly IPython Notebooks) is a web-based interactive web application for creating documents that contain live code, equations, visualizations and narrative text. While notebooks represent a good way to build a tutorial or a report, you will still need to rely on more common, scripting IDEs like Spyder or PyCharm, when working on a "project" (e.g. by organizing functions in modules to industrialize a data science product).

The notebook consists of a sequence of cells. A cell is a multiline text input field, and **its contents can be executed** by using `Shift+Enter`, or by clicking either the `Play` button on the toolbar, or Cell > Run in the menu bar. The execution behavior of a cell is determined by the cell’s type. There are three types of cells: *code cells*, *markdown cells*, and *raw cells*. Every cell starts off being a code cell, but its type can be changed by using a drop-down on the toolbar (which will be “Code”, initially). **As part of this course, DO NOT change the type of a pre-existing cell, DO NOT duplicate a pre-existing cell**. However, you can create and change the type of any new cell.

## Code cells

A code cell allows you to edit and write new code, with full syntax highlighting and tab completion. The programming language used here is Python.

When a code cell is executed, code that it contains is sent to the kernel associated with the notebook. The results that are returned from this computation are then displayed in the notebook as the cell’s output. The output is not limited to text, with many other possible forms of output are also possible, including figures and HTML tables.

**Tips**:
- press the `Tab` key to use auto completion in code cells
- press `Shift+Tab` to display the documentation of the current object (the one on witch the cursor is)

See e.g. [this introduction](https://miykael.github.io/nipype_tutorial/notebooks/introduction_jupyter-notebook.html).

## Markdown cells

You can document the computational process in a literate way, alternating descriptive text with code, using rich text. In IPython this is accomplished by marking up text with the Markdown language. The corresponding cells are called Markdown cells. The Markdown language provides a simple way to perform this text markup, that is, to specify which parts of the text should be emphasized (italics), bold, form lists, etc. See [this cheatsheet](https://www.edureka.co/blog/wp-content/uploads/2018/10/Jupyter_Notebook_CheatSheet_Edureka.pdf).

Markdown cells will also render $\LaTeX$ code escaped with \\$\LaTeX\\$.

For more information about Jupyter Notebooks, read its [official documentation](https://jupyter-notebook.readthedocs.io/en/stable/notebook.html#structure-of-a-notebook-document).

## Read-only cells

Most cells in these assignments have been marked "read-only" meaning you will not be able to modify their content, to avoid mistakes and guarantee some reproducibility when grading them. If you need to inspect what is going on inside such a cell, create a new cell and copy over the content that you want to run/inspect.

## `NotImplementedError`

Whenever you see a (modifiable) cell with a `raise NotImplementedError` instruction, it means you should implement something! This cell will generally be preceded by a Markdown cell explaining what you need to do. The following cell usually contains visible and/or hidden tests (to test your code). If you see (seemingly) empty cells after a code cell where you implemented something, it usually means there are only hidden tests that are run against your code! 

## Save your work!

Don't forget to regularly save your work! There is some autosaving but hitting `Ctrl+S` will ease your mind.

## Validate button

If you correctly followed the [guidelines to setup your Python environment](https://moodle.ip-paris.fr/mod/page/view.php?id=180914), you should see a "Validate" button in the tool bar, next to the Save, +, Scissors, ...
Hitting this button, it should turn grey and say "Validating..." and show some message at some later point, e.g. "Success! Your notebook
passes all the tests!". Please use and abuse this button before submitting your assignment on Moodle.

## Objectives

What we will practice today:
- Pandas basics
- Making an exploratory analysis on an real dataset
- Making a first predictive model for this datase

# Basic built-in tools

Contrary to C, C++, or Java, Python is **not** statically typed, i.e. you can do:

```python
a = 1
a = "bob"
a
```

Nevertheless, in this course, whenever possible, we will give type **hints**, such as:

```python
from typing import List

def len(l: List[int]):
    cnt = 0
    for _ in l:  # "_" means "discard current value"
        cnt += 1
    return cnt
```

This doesn't enforce anything, i.e. you can call `len` on a list of strings,
but it helps with readability and, most importantly, helps you understand what is expected
from the function (and you, implementing it).

```python
print(len([1, 2]))
len(["alice", "bob"])
```

## Loops

Python has two primitive loop commands, `while` and `for`.
They can be used for iterating over a sequence (that is either a list, a tuple, a dictionary, a set, a string,
an iterator).
Looping through lists, tuples, dictionaries or sets gives you one element contained in it until exhaustion
(we'll see later on how).
Similarly, looping through a string will give you each character.

Loops are associated with two special keywords: `continue` (i.e. prematurely do the next iteration) and `break`
(prematurely get out of the loop).

```python
from typing import List

fruits = ["apple", "banana", "cherry"]


def find_banana(l: List[str]):
    for x in l:
        if x == "banana":
            print("found")
            break
        if x == "cherry":
            continue
        print("neither a banana nor a cherry")


find_banana(fruits)
```

## Function definition, code organization

As can be seen from the earlier examples, a function can be defined with the `def`
keyword, followed by its name, parentheses possibly containing arguments and/or
keyword arguments (`*args` and `**kwargs` are generic "holders" for these).

```python
def my_func(a, *args, b: int = 3, **kwargs):
    print(f"a: {a}")
    print(f"args: {args}")
    print(f"b: {b}")
    print(f"kwargs: {kwargs}")


my_func(1, 'a', 'b', 'c', b=2, c=4, d="toto")
```

## Data structures

Data structures or containers often refer to objects capable of
holding several other objects (of the same or of various types)
and perform common operations on it.

The most common such structure in Python is the `list`, capable
of holding objects of any type and providing useful features
like `pop(index)` (returning the index-th element and deleting
it from the list), `append(element)` (adding an element at the
end), `extend(l)` (to concatenate it in place with another list).

Another common structure in Python is the `dict`ionary. It holds
key--value pairs of any types.

A very Pythonic trick is to use `list`- (resp. `dict`-)
comprehensions, that is, looping through each element
inside an expression that returns a `list` (resp. `dict`):

```python
l = [1, 2, 3]
[el + 1 for el in l]
```

## Conventions and style

As in many programming languages, people have come up
with many ways of using it. In this course, we will try
to conform to [PEP-styling](https://peps.python.org/pep-0008/):
a set of style recommendations
that many developers and organizations follow.

In particular, modules, functions and attributes will use
snake_case formatting while classes use CamelCase.

## Object-Oriented Programming (OOP)

In this course, whenever useful and possible, we will use the
OOP paradigm. It is particularly used in Python since
[everything is an object, even classes](https://linux.die.net/diveintopython/html/getting_to_know_python/everything_is_an_object.html#d0e4665).

### What's a class? Encapsulation.

A class is usually appropriate when it is possible
to _encapsulate_ or _bundle_ related information
together in some meaningful way:

```python
class Person(object):
    def __init__(self, first_name: str, last_name: str):
        self.first_name = first_name
        self.last_name = last_name


person = Person("Adrien", "Ehrhardt")
print(person.first_name)
```

### Magic methods: abstraction.

Once information is encapsulated in some meaningful
container, you can abstract away some logical operation(s)
on the objects by hiding how this happens in _class methods_.
Python comes with _magic methods_ that are "standard" class
methods' names and will be automatically called when doing
some operations.
We already saw the `__init__` method that gets called when
instantiating a new object. Some other useful magic methods:

```python
class Person(object):
    def __init__(self, first_name: str, last_name: str):
        print("init")
        self.first_name = first_name
        self.last_name = last_name

    def __new__(cls, first_name: str, last_name: str):
        """Creates a new object and returns it, effectively passing it to __init__"""
        print("new")
        new_obj = super().__new__(cls)
        return new_obj

    def __call__(self):
        print(f"Hey {self.first_name}!")

    def __str__(self):
        return f"{self.first_name} {self.last_name}"

    def __repr__(self):
        return f"{id(self)} - {self.first_name} - {self.last_name}"

person = Person("Adrien", "Ehrhardt")
print(repr(person))
person()
print(f"Person person is {person}")
```

Apart from object creation, calling and identification, you may want to add operators, such as
`__len__` (which is called with `len(object)`), `__add__` (`obj1 + obj2`), `__gt__` (`>`), etc.

Lastly (for this course), iterators can be made with the `yield` keyword. It represents a countable number of values
that can be iterated over. Lists, tuples, dictionaries, and sets are all iterable objects, meaning all the objects
they contain are already _materialized_ (in particular, you can take the element at index $i$ of a list of size greater
than $i$). In contrast, iterators are objects of classes that can `yield` the next element to iterate over "as you go".

```python
def my_iterator(length: int = 10):
    cur_length = 0
    while cur_length < length:
        yield (cur_length := (cur_length + 1))


for i in my_iterator():
    print(i)
```

At any point in time, the past values iterated over are discarded, and the future values are unknown, which saves
a lot of time and memory:

```python
from timeit import timeit

print(timeit(lambda: range(100), number=100))
print(timeit(lambda: my_iterator(100), number=100))
print(timeit(lambda: list(range(100)), number=100))
```

### A look at inheritance and polymorphism through base/abstract classes

Let's take the earlier `Person` example yet again, and assume that any person must have a job which gives them
additional properties. This would be nicely implemented with an abstract class:

```python
from abc import ABC, abstractmethod

class Person(ABC):  # ABC means Person is now abstract
    def __init__(self, first_name: str, last_name: str):
        self.first_name = first_name
        self.last_name = last_name

    @abstractmethod  # any non-abstract inheriting class must implement this
    def occupation(self):
        pass

    def __str__(self):
      return f"{self.first_name} {self.last_name} {self.occupation()}"
    

class Teacher(Person):
    def occupation(self):
        return "teaches INF442"


print(Teacher("Adrien", "Ehrhardt"))
Person("Adrien", "Ehrhardt")
```

# NumPy

One of Python's strengths is its wide adoption. With wide adoption
(usually) comes wide availability of third-party **libraries** (yet
another name for a Python package that has been distributed), s.t.
you don't need to code matrix multiplication each time you use Python!

One such library is NumPy, which is actually written in C and is only some
kind of interface between you (well, through Python) and operations that are implemented in C.
Indeed, a number of "critical" things (things that must happen fast) in
Python are implemented in "fast"/low-level languages and what you
get is just some _API_ to interact with that lower-level library.

On NumPy specifically, these resources can be useful to the interested reader:
[Python programming for data science](https://www.tomasbeuzen.com/python-programming-for-data-science/chapters/chapter5-numpy.html) and
[Python programming for data science addendum](https://www.tomasbeuzen.com/python-programming-for-data-science/chapters/chapter6-numpy-addendum.html).

NumPy provides vectors/matrices of arbitrary dimensions and implements some useful operations.

```python
import numpy as np

print(np.zeros(3))
print(np.ones((2, 3)))
```

A NumPy array is relatively low-level, in the sense that it is, e.g., singly typed.

1 will be cast to string:

```python
np.array(['a', 1])
```

Adding arrays with incompatible dtypes will fail:

```python
np.ones(2) + np.array(['a', 'b'])
```

The underlying type can be retrieved:

```python
np.ones(2).dtype
```

Although low-level, it abstracts many very useful operations.
The documentation provides many details, but some (if not all) useful NumPy commands for this course are:

- `np.zeros`, `np.ones` which take a shape as argument and return an array of the requested shape filled
  with zeros (resp. ones).

- `np.full` to achieve the same behavior and provide an arbitrary `fill_value`.

- `np.transpose(array)`, or `array.T` to transpose an array.

- `np.(nan)(arg){min,max}(array)` or `array.{min,max}()` to get the (arg){min,max} of an array, possibly
  discarding any null (nan) values encountered.

- `array[:, 1]` to **slice** the 2nd column and all rows.

Recommendation: always prefer "vector(ized)" operations (i.e. `array[:, 1] = 0`)
over loops (i.e. `for row in range(len(array)): A[row, 1] = 0`).

NumPy allows **broadcasting** (which already took place in the earlier example: `0` was "interpreted" as vector array
of 0s on the right hand side of `array[:, 1] = 0`).
While this topic is out of the scope of the present course
(the interested reader can go to [this page](https://numpy.org/doc/stable/user/basics.broadcasting.html)
for more details), you will much likely encounter errors like this:

```text
ValueError: operands could not be broadcast together.
```
Although low-level, it abstracts many very useful operations.
The documentation provides many details, but some (if not all) useful NumPy commands for this course are:

- `np.zeros`, `np.ones` which take a shape as argument and return an array of the requested shape filled
  with zeros (resp. ones).

- `np.full` to achieve the same behavior and provide an arbitrary `fill_value`.

- `np.transpose(array)`, or `array.T` to transpose an array.

- `np.(nan)(arg){min,max}(array)` or `array.{min,max}()` to get the (arg){min,max} of an array, possibly
  discarding any null (nan) values encountered.

- `array[:, 1]` to **slice** the 2nd column and all rows.

Recommendation: always prefer "vector(ized)" operations (i.e. `array[:, 1] = 0`)
over loops (i.e. `for row in range(len(array)): A[row, 1] = 0`).

NumPy allows **broadcasting** (which already took place in the earlier example: `0` was "interpreted" as vector array
of 0s on the right hand side of `array[:, 1] = 0`).
While this topic is out of the scope of the present course
(the interested reader can go to [this page](https://numpy.org/doc/stable/user/basics.broadcasting.html)
for more details), you will much likely encounter errors like this:

```text
ValueError: operands could not be broadcast together.
```

This usually means that you're trying to perform on array operation implying arrays of incompatible shapes,
e.g. multiplying a (3, 4) by a (2, 5) matrix.

```python
np.ones((3, 4)) @ np.ones((2, 5))
```

# Pandas Basics

[Pandas](http://pandas.pydata.org/) is a popular data analysis toolkit for Python.

We will use it to explore data with heterogeneous types and/or missing values.

Additional documentation can be found there:
- https://pandas.pydata.org/docs/getting_started/index.html
- http://pandas.pydata.org/pandas-docs/stable/
- https://jakevdp.github.io/PythonDataScienceHandbook/
- http://www.jdhp.org/docs/notebook/python_pandas_en.html

## Import directives

To begin with, let's import the Pandas library as the *pd* alias. Select the following code cell and execute it with the `Shift + Enter` shortcut key.

We also import:
- Numpy to generate arrays
- Seaborn to make plots
- Sklearn to learn data
- Matplotlib to plot some stuff

```python
import numpy as np
import pandas as pd
import seaborn as sns
import matplotlib.pyplot as plt
import sklearn
```

## Make DataFrames (2D data)

Let's make a *DataFrame* now. DataFrame is the name given to 2D arrays in Pandas. We will assign this DataFrame to the `df` variable and display it (with the `df` line at the end of the cell).

When the data carries a meaning, i.e. comes from a real data source, e.g. the `titanic` dataset that we will load later on consisting in characteristics of passengers of the Titanic, always use explicit names such as `df_titanic`. See [PEP8](https://www.python.org/dev/peps/pep-0008/), a Python code style guide, for details.

```python
data = [[1, 2, 3], [4, 5, 6]]
df = pd.DataFrame(data)
df
```

The previous command has made a DataFrame with automatic *indices* (rows index) and *columns* (columns label).

To make a DataFrame with a custom index (`pandas` allows to define multiple indices gathered in a `MultiIndex` but this feature will not covered in this course) and custom column names, use:

```python
data = [[1, 2, 3], [4, 5, 6]]
index = [10, 20]
columns = ['A', 'B', 'C']

df = pd.DataFrame(data=data, index=index, columns=columns)
df
```

A Python dictionary can be used to define data; its keys define its **columns'** labels.

```python
data_dict = {'A': 'foo',
             'B': [10, 20, 30],
             'C': 3.14}
df = pd.DataFrame(data_dict, index=[10, 20, 30])
df
```

### Get information about a dataframe

Its (row) index (possibly a `MultiIndex`):

```python
df.index
```

Its columns' labels (also an index!):

```python
df.columns
```

Its shape (i.e. number of rows and number of columns):

```python
df.shape
```

The number of lines in `df` is:

```python
df.shape[0]
```

Or equivalently:

```python
len(df)
```

The number of columns in `df` is:

```python
df.shape[1]
```

The data type of each column:

```python
df.dtypes
```

Additional information about columns:

```python
df.info()
```

```python
df.describe()
```

### Select a single column

Here are 3 equivalent syntaxes to get the column "C":

```python
df.C
```

```python
df["C"]
```

```python
df.loc[:, "C"]
```

Note that a "column" corresponds to a `pandas.Series` object:

```python
type(df["C"])
```

Selecting multiple columns, or a list containing at least one column, will output a `pandas.DataFrame`. This is subtle but may be the cause of subsequent errors.

```python
df[["C"]]
```

```python
type(df[["C"]])
```

### Select multiple columns

Here are 2 equivalent syntaxes to get both columns "A" and "B":

```python
df[['A', 'B']]
```

```python
df.loc[:, ['A', 'B']]
```

### Select a single row

```python
df.loc[10]
```

```python
df.loc[10, :]
```

Note that again, the result being a "vector", the output is of type `pandas.Series`:

```python
type(df.loc[10, :])
```

### Select multiple rows

```python
df.loc[[10, 20], :]
```

### Select rows based on values

```python
df.B < 30.  # this is a Series
```

```python
df.loc[df.B < 30.]  # this is a DataFrame
```

```python
df.loc[(df.B < 20) | (df.B >= 30)]
```

```python
df.loc[(df.B >= 20) & (df.B < 30)]
```

### Select rows and columns

```python
df.loc[(df.B < 20) | (df.B >= 30), 'C']  # this is a Series
```

```python
df.loc[(df.B < 20) | (df.B >= 30), ['A', 'B']]  # this is a DataFrame
```

### Select columns based on a condition

Which column of the 10th line (starting from the Oth) takes value `'foo'`?

```python
df.loc[10] == 'foo'
```

Of those, which value(s) are `True`?

**The 0th**

```python
np.where(df.loc[10] == 'foo')
```

To which column name (or `Index`) does column 0 correspond?

```python
df.columns[np.where(df.loc[10] == 'foo')]
```

```python
df.columns[np.where(df.loc[10] == 'foo')].to_list()
```

### Apply a function to selected columns' values

```python
df.B *= 2.
df
```

```python
df.B = pow(df.B, 2)
df
```

### Apply a function to selected rows' values

```python
df.loc[df.B < 500., 'A'] = -1
df
```

```python
df.loc[(df.B < 500.) | (df.B > 2000), 'C'] = 0
df
```

### Handling missing data

Missing data are represented by a *NaN* ("Not a Number").

```python
# equivalently, use np.nan
data = [[3, 2, 3],
        [float("nan"), 4, 4],
        [5, float("nan"), 5],
        [float("nan"), 3, 6],
        [7, 1, 1]]
df = pd.DataFrame(data, columns=['A', 'B', 'C'])
df
```

To obtain the boolean mask where values are NaN, type:

```python
df.isnull()
```

To drop any rows that have missing data:

```python
df.dropna()
```

Note: this does not replace df (as suggested by the fact that a DataFrame is printed: something is returned!).

```python
df
```

If you wish to overwrite `df`, you may reassign (`df = ...`) or use `inplace=True` (common to lots of `pandas`' methods).

To fill missing data with a chosen value (e.g. 999):

```python
df.fillna(value=999)  # Again, this returns a new DataFrame
```

To count the number of NaN values in a given column:

```python
df.isnull().sum()
```

The `axis` keyword (common to lots of `pandas`' methods) allows to do the summation on each line rather than by column:

```python
df.isnull().sum(axis=1)
```

### Exercice 1

```python
df = pd.DataFrame([[i + j * 10 for i in range(10)] for j in range(20)],
                  columns=["A", "B", "C", "D", "E", "F", "G", "H", "I", "J"])
df
```

Considering the dataframe `df` created in the previous cell:
- Write the code to extract the column 'C' and assign this subarray to a variable `df2` **of type DataFrame**.

<!-- grade_id: cell-439df74f7e4ad4f6, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-dd2e546b29de2a17, points: 1, grade: true, solution: false -->
```python
assert type(df2) == pd.DataFrame
```

- Write the code to extract the column 'A' and 'F' and assign this sub-`DataFrame` to a variable `df3`.

<!-- grade_id: cell-3607eab496454b0a, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-78de20a754644993, points: 1, grade: true, solution: false -->
```python
assert type(df3) == pd.DataFrame
```

- Write the code to extract lines where the value in column 'D' is greater than 100 and assign this sub-`DataFrame` to a variable `df4`.

<!-- grade_id: cell-7213224e6e999a5e, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-81b38dbd7701e36a, points: 1, grade: true, solution: false -->
```python
assert type(df4) == pd.DataFrame
```

- Write the code to extract lines where the value in column 'D' is greater than 100 or less than 40 and assign this sub-`DataFrame` to a variable `df5`.

<!-- grade_id: cell-56475e3940244b7e, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-3137844c56627d82, points: 1, grade: true, solution: false -->
```python
assert type(df5) == pd.DataFrame
```

- Write the code to extract lines where the value in column 'D' is less than 100 and greater than 40 and assign this sub-`DataFrame` to a variable `df6`.

<!-- grade_id: cell-a6afe8f7c217ee0b, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-cb5b5ea5c60ff474, points: 1, grade: true, solution: false -->
```python
assert type(df6) == pd.DataFrame
```

- Write the code to extract the values of columns 'D' and 'E' of lines where column 'B' is greater than 100 or less than 40 and assign this sub-`DataFrame` to a variable `df7`.

<!-- grade_id: cell-0238b0c446614524, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-d33dde24d7b3c076, points: 1, grade: true, solution: false -->
```python
assert type(df7) == pd.DataFrame
```

**N.B.:** some companies / labs are shifting from pandas to [polars](https://pola.rs/) which will **NOT** be covered in this course. Its main advantages are speed and lazy executio (but there's no index). [Read more](https://docs.pola.rs/user-guide/migration/pandas/)

# Exploratory analysis of the Titanic dataset

## Problem description

The sinking of the RMS Titanic is one of the most infamous shipwrecks in history.  On April 15, 1912, during her maiden voyage, the Titanic sank after colliding with an iceberg, killing 1502 out of 2224 passengers and crew. This sensational tragedy shocked the international community and led to better safety regulations for ships.

One of the reasons that the shipwreck led to such loss of life was that there were not enough lifeboats for the passengers and crew. Although there was some element of luck involved in surviving the sinking, some groups of people were more likely to survive than others, such as women, children, and the upper-class.

In this lab session, you will complete the analysis of what sorts of people were likely to survive and you will apply machine learning methods to predict which passengers survived the tragedy.

([Description from Kaggle](https://www.kaggle.com/c/titanic))

## Load data

We start by acquiring the dataset into the `df` Pandas DataFrames.

```python
titanic = pd.read_csv("https://raw.githubusercontent.com/ogrisel/parallel_ml_tutorial/master/notebooks/titanic_train.csv")
titanic
```

## Variables description

Here are the *features* available in the dataset:

```python
for column in titanic.columns:
    print(column)
```

- *Survived*: survived (1) or died (0)
- *Pclass*: passenger's class
- *Name*: passenger's name
- *Sex*: passenger's sex
- *Age*: passenger's age
- *SibSp*: number of siblings/spouses aboard
- *Parch*: number of parents/children aboard
- *Ticket*: ticket number
- *Fare*: fare
- *Cabin*: cabin
- *Embarked*: port of embarkation
  - C = Cherbourg
  - Q = Queenstown
  - S = Southampton


([Description from Kaggle](https://www.kaggle.com/c/titanic/data))

## Exercise 2

- List *categorical* features in the Titanic dataset and store it in `titanic_categorical` (of type `list`).

**Hint**: the types of all features can be retrieved with `titanic.dtypes`. Which one(s) are **not** *numerical*?

**Warning**: a systematic approach is expected, DO NOT simply copy/paste column names from above.

<!-- grade_id: cell-bb908cf4f12ae20e, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-b2cfd1859c56dece, points: 1, grade: true, solution: false -->
```python

```

- List *numerical* features and store it in `titanic_numerical`. **Hint**: if they're numerical, they're not *categorical*.

<!-- grade_id: cell-31eec71b64072565, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-c3f19e3115114386, points: 0, grade: true, solution: false -->
```python

```

- Which features contain blank, null or empty values? Store a list of them in `titanic_features_blank`.

**Warning**: a list of feature names is expected.

<!-- grade_id: cell-760b953ec1c8855b, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-6d38b35ec18f84c1, points: 1, grade: true, solution: false -->
```python

```

- We want to complete the analysis of what sorts of people were likely to survive, i.e. to predict for a new passenger whether he/she will (or would have...) survive(d) or not. What kind of machine learning problem is it?
- A supervised / unsupervised problem?
- A classification / regression / clustering / reinforcement learning problem?
- **Why?**

<!-- grade_id: cell-1c1a246a13908fd9, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

## Display a quick summary of the dataset

```python
titanic.describe()
```

## Exercise 3

- What is the age of the oldest passenger (or crew)? Store it in `age_oldest_on_titanic`.

**Warning**: use a systematic approach. DO NOT simply copy/paste from the table above. A `float` is expected.

<!-- grade_id: cell-8426f2d06daab9a9, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-b84657a5d8e82d12, points: 1, grade: true, solution: false -->
```python

```

- What is the average fare ? Store it in `average_fare_on_titanic`.

**Warning**: use a systematic approach. DO NOT simply copy/paste from the table above. A `float` is expected.

<!-- grade_id: cell-averageFARE1, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-3d92aacb2c30d921, points: 1, grade: true, solution: false -->
```python

```

- What are the fare quartiles (the 25%, 50%, 75% quantiles)?

**Warning**: use a systematic approach. DO NOT simply copy/paste from the table above. A `Series` is expected.

<!-- grade_id: cell-f8af80abc411bfc5, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

## Explore correlations between the numerical features and survival of passengers

Let us start by understanding correlations between numerical features and the label we want to predict (Survived).

### Correlation survival vs age

```python
bins = np.arange(0, 100, 5)  # this controls the width of the age categories
print("Bins:", bins)

ax = titanic.loc[titanic.Survived == 0, "Age"].hist(  # we select dead people's rows and use hist(ogram)
    bins=bins, color="red", alpha=0.5, label="died", figsize=(18, 8))  # let's plot them in red
titanic.loc[titanic.Survived == 1, "Age"].hist(  # we select surviving people's rows and use hist(ogram)
    bins=bins, color="blue", ax=ax, alpha=0.5, label="survived")  # let's plot them in blue
ax.legend();  # let's add a legend
plt.show()
```

Don't forget to check missing values as it could bias the plot above!

```python
titanic.loc[(titanic.Age.isnull()) & (titanic.Survived == 0)].shape[0]
```

```python
titanic.loc[(titanic.Age.isnull()) & (titanic.Survived == 1)].shape[0]
```

## Exercise 4

What useful qualitative observation can you extract from the previous plot?

<!-- grade_id: cell-a41bfe53c8227bc9, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

## Exercise 5

Plot correlation between Fare and Survival. Use the code from the previous cell as a starting point.

<!-- grade_id: cell-f1c81ab5ad32fc5b, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

What useful observation can you extract from the previous plot?

<!-- grade_id: cell-93819835ee94b1c3, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

## Explore correlations between the categorical features and survival of passengers

### Correlation survival vs class

```python
g = sns.catplot(y="Survived", x="Pclass", kind="bar", data=titanic)
g.set_ylabels("survival probability");
```

Survival probability is simply the ratio of surviving passengers in each class.

## Exercise 6

What useful observation can you extract from the previous plot?

<!-- grade_id: cell-61418c02363f3250, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

Plot correlation between
* Sex and Survival
* SibSp and Survival
* Parch and Survival
* Embarked and Survival

Use the code from the previous cell as a starting point.

<!-- grade_id: cell-57afbe699faf2a84, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

What useful observation can you extract from these plots ?

<!-- grade_id: cell-dce697156e2e038b, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

## Exercise 7

The following code computes the survival rate of women in the first class, i.e. an estimate of $P(\text{Survived} = 1 | \text{Sex}=\text{female}, \text{Pclass}=1)$.

Note how we're making good use of the (frustrating) first exercises.

```python
p_survived_1_given_female_1 = titanic.loc[(titanic.Sex == "female") & (titanic.Pclass == 1), "Survived"].mean()
p_survived_1_given_female_1
```

Compute the survival rate corresponding to:

- $P(\text{Survived} = 1 | \text{Sex}=\text{female}, \text{Pclass}=2)$ and store it in `p_survived_1_given_female_2`

<!-- grade_id: cell-9927db2d2f26d325, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
p_survived_1_given_female_2
```

<!-- grade_id: cell-264b06d3d0b6cbb4, points: 1, grade: true, solution: false -->
```python

```

- $P(\text{Survived} = 1 | \text{Sex}=\text{female}, \text{Pclass}=3)$ and store it in `p_survived_1_given_female_3`

<!-- grade_id: cell-4883eca7937bebe7, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
p_survived_1_given_female_3
```

<!-- grade_id: cell-5b2947a657be54e4, points: 1, grade: true, solution: false -->
```python

```

- $P(\text{Survived}=1 | \text{Sex}=\text{male}, \text{Pclass}=1)$ and store it in `p_survived_1_given_male_1`

<!-- grade_id: cell-f524d6a84a768d49, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
p_survived_1_given_male_1
```

<!-- grade_id: cell-374cc7302701a232, points: 1, grade: true, solution: false -->
```python

```

- $P(\text{Survived} = 1 | \text{Sex}=\text{male}, \text{Pclass}=2)$ and store it in `p_survived_1_given_male_2`

<!-- grade_id: cell-eff206b0fa159947, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
p_survived_1_given_male_2
```

<!-- grade_id: cell-d26e9e2da5e80004, points: 1, grade: true, solution: false -->
```python

```

- $P(\text{Survived}=1 | \text{Sex}=\text{male}, \text{Pclass}=3)$ and store it in `p_survived_1_given_male_3`

<!-- grade_id: cell-de3af50bfb0c175b, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
p_survived_1_given_male_3
```

<!-- grade_id: cell-4e260dc4624abb54, points: 1, grade: true, solution: false -->
```python

```

The following plots display the survival rate considering multiple variables. This might exhibit "multivariate" or "compound" effects on survival.

```python
sns.catplot(x="Pclass", y="Survived", hue="Sex", kind="bar", data=titanic)
g.set_ylabels("survival probability");
```

```python
sns.catplot(y="Survived", x="Embarked", hue="Pclass", kind="bar", data=titanic)
g.set_ylabels("survival probability");
```

Explore other combinations of features with similar plots.

After considering these plots, which variables do you expect to be good predictors of survival?

<!-- grade_id: cell-162f2ce12f82a874, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

# Make a predictive model

After this brief exploration of the dataset, let's try to train a model to predict the survival of "new" (i.e. unknown) passengers.

A *decision tree* classifier will be used to complete this task.

To begin with, import the decision tree package (named `tree`) implemented in Scikit Learn (a.k.a. `sklearn`).

```python
import sklearn.tree
```

Then reload the dataset:

```python
titanic = pd.read_csv("https://raw.githubusercontent.com/ogrisel/parallel_ml_tutorial/master/notebooks/titanic_train.csv")
```

## Exercise 8

Based on investigations made in the exploratory analysis, which variables can be ignored (they cannot convey any meaningful information - think IDs...)? Complete the following cell to remove useless features.

<!-- grade_id: cell-e2bbe6a5c9c2f820, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-c4d608f0eab25512, points: 1, grade: true, solution: false -->
```python

```

Once you have removed useless variables, remove examples with missing values from the dataset.

<!-- grade_id: cell-a7cf35b9f8fb6849, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-d232b59289affc54, points: 1, grade: true, solution: false -->
```python

```

Let's finish the dataset cleaning by converting categorical features to numerical ones. This is often required by Machine Learning algorithms' implementations (in particular in `sklearn`) since most os these algorithms rely on matrix calculus / algebra as we will see later on in this course.

```python
titanic['Embarked'] = titanic['Embarked'].map({'S': 0, 'C': 1, 'Q': 2}).astype(int)
titanic['Sex'] = titanic['Sex'].map({'male': 0, 'female': 1}).astype(int)
```

The following cell gives an overview of our dataset (only the first lines displayed).

```python
titanic.head()
```

Now that the dataset is ready, we split it in two subsets:
- the training set (`X_train` and `Y_train`);
- the testing set (`X_test` and `Y_test`).

`X_xxx` contains example's *features* and `Y_xxx` contains example's *labels*.

```python
X = titanic.drop("Survived", axis=1)  # X cannot contain the label
Y = titanic["Survived"]

X_train = X.iloc[:-10]  # all samples except 10 last ones
Y_train = Y.iloc[:-10]

X_test = X.iloc[-10:]  # 10 last samples
Y_test = Y.iloc[-10:]

X_train.shape, Y_train.shape, X_test.shape, Y_test.shape
```

## Exercise 9

Explain qualitatively your intuition about why it is necessary to split the dataset in these two subsets (training set - think **past** results of dice rolls {1, 2, 1} - and testing set - think **new** dice rolls {6, 4}) to understand the behavior / performance a Machine Learning model (think predicting the average of a dice throw).

<!-- grade_id: cell-53ca2e22509dfdcb, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

The classifier is instanciated (it is now an object of a class) and trained (thanks to a method implemented in this class) with the following code:

```python
decision_tree = sklearn.tree.DecisionTreeClassifier()

decision_tree.fit(X_train, Y_train)
```

Note that in `sklearn` most algorithms are implemented as classes.

A particular instance is created, and we use class methods such as `fit`, `transform` or `predict` on appropriate arguments to perform operations.

Then the success rate of predictions is checked on both sets:

```python
decision_tree.score(X_train, Y_train)
```

```python
decision_tree.score(X_test, Y_test)
```

Finally we make some predictions and compare them to the truth:

```python
Y_pred = decision_tree.predict(X_test)

pd.DataFrame(np.array([Y_pred, Y_test]).T, columns=('Predicted', 'Actual'))
```

## Plot feature importances (for the trained model)

The following code plot the relative importance of features for the prediction task (for the trained model).

(What "importance" means is beyond the scope of this lab; the interested reader might refer to `sklearn`'s documentation)

```python
imp = pd.DataFrame(decision_tree.feature_importances_,
                   index=X_train.columns,
                   columns=['Importance'])
imp.sort_values(['Importance'], ascending=False)
```

```python
imp.sort_values(['Importance'], ascending=False).plot.bar();
```

## Exercise 10

Compare this list to the assumptions made during the exploratory analysis.
Have you predicted a similar ranking?

<!-- grade_id: cell-d8bae2a924185040, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

## Display the decision tree

```python
from sklearn import tree
tree.plot_tree(decision_tree);
```

Something should be striking and feel wrong on this plot... Try to think about what could have happened!

# Going further

Here is a good tutorial to complete the lab session: [Understanding and diagnosing your machine-learning models](http://gael-varoquaux.info/interpreting_ml_tuto/) (by Gael Varoquaux)